---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Launch graph and fit the line

```yaml
type: VideoExercise
key: 7e7b2599d3
xp: 50
```

`@projector_key`
b5e84ac8dfe63b0ee9f2b2fd35cb3420
